<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9: ANOVA table - Millet</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>
[[lang code="en"]]An experiment is conducted on {@plots@} plots to understand the effect of several factors on the yield of millet.[[/lang]]
[[lang code="fr"]]Une expérience est menée sur {@plots@} parcelles afin d’étudier l’effet de plusieurs facteurs sur le rendement du mil.[[/lang]]
</p>

<p>
[[lang code="en"]]The factors are: {@variety@} varieties of millet, {@weeding@} types of weeding and {@fert@} different levels of fertiliser.[[/lang]]
[[lang code="fr"]]Les facteurs sont : {@variety@} variétés de mil, {@weeding@} types de désherbage et {@fert@} différents niveaux de fertilisation.[[/lang]]
</p>

<p>
[[lang code="en"]]Complete the ANOVA table for the experiment.[[/lang]]
[[lang code="fr"]]Complétez le tableau d’ANOVA de l’expérience.[[/lang]]
</p>

<p>
<b>
[[lang code="en"]]Give the values to \(1\) decimal place.[[/lang]]
[[lang code="fr"]]Donnez les valeurs à \(1\) décimale.[[/lang]]
</b>
</p>

<table>
    <caption></caption>
    <thead>
        <tr>
            <th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Term[[/lang]][[lang code="fr"]]Terme[[/lang]]</th>
            <th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Df[[/lang]][[lang code="fr"]]ddl[[/lang]]</th>
            <th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Sum Sq[[/lang]][[lang code="fr"]]Somme des carrés[[/lang]]</th>
            <th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Mean Sq[[/lang]][[lang code="fr"]]Carré moyen[[/lang]]</th>
            <th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]F Value[[/lang]][[lang code="fr"]]Valeur de F[[/lang]]</th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Variety[[/lang]][[lang code="fr"]]Variété[[/lang]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_variety_df]] [[validation:ans_variety_df]] [[feedback:prt1_variety_df]]</td>
            <td style="text-align: center;border: 1px solid black;">{@variety_ss@}</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_variety_ms]] [[validation:ans_variety_ms]] [[feedback:prt_variety_ms]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_variety_f]] [[validation:ans_variety_f]] [[feedback:prt_variety_f]]</td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Weeding[[/lang]][[lang code="fr"]]Désherbage[[/lang]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_weeding_df]] [[validation:ans_weeding_df]] [[feedback:prt_weeding_df]]</td>
            <td style="text-align: center;border: 1px solid black;">{@weeding_ss@}</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_weeding_ms]] [[validation:ans_weeding_ms]] [[feedback:prt_weeding_ms]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_weeding_f]] [[validation:ans_weeding_f]] [[feedback:prt_weeding_f]]</td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Fertiliser[[/lang]][[lang code="fr"]]Engrais[[/lang]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_fert_df]] [[validation:ans_fert_df]] [[feedback:prt_fert_df]]</td>
            <td style="text-align: center;border: 1px solid black;">{@fert_ss@}</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_fert_ms]] [[validation:ans_fert_ms]] [[feedback:prt_fert_ms]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_fert_f]] [[validation:ans_fert_f]] [[feedback:prt_fert_f]]</td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Residuals[[/lang]][[lang code="fr"]]Résidus[[/lang]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_resid_df]] [[validation:ans_resid_df]] [[feedback:prt_resid_df]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_resid_ss]] [[validation:ans_resid_ss]] [[feedback:prt_resid_ss]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_resid_ms]] [[validation:ans_resid_ms]] [[feedback:prt_resid_ms]]</td>
            <td style="text-align: center;border: 1px solid black;"></td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Total[[/lang]][[lang code="fr"]]Total[[/lang]]</td>
            <td style="text-align: center;border: 1px solid black;">[[input:ans_total_df]] [[validation:ans_total_df]] [[feedback:prt1_total_df]]</td>
            <td style="text-align: center;border: 1px solid black;">{@total_ss@}</td>
            <td style="text-align: center;border: 1px solid black;"></td>
            <td style="text-align: center;border: 1px solid black;"></td>
        </tr>
    </tbody>
</table>

<p></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>
[[lang code="en"]]The complete ANOVA table is shown below.[[/lang]]
[[lang code="fr"]]Le tableau d’ANOVA complet est présenté ci-dessous.[[/lang]]
</p>

<p>
[[lang code="en"]]There are a few key points to remember:[[/lang]]
[[lang code="fr"]]Quelques points essentiels sont à retenir :[[/lang]]
</p>

<ul>
    <li>
        [[lang code="en"]]The residual is the "left over" after taking into account the variables. So the residual Df and sum of squares is the difference between the total values and the sum of the variables values.[[/lang]]
        [[lang code="fr"]]Le résidu correspond à ce qui reste après avoir pris en compte les variables. Ainsi, les degrés de liberté et la somme des carrés résiduels correspondent à la différence entre les valeurs totales et la somme des valeurs associées aux variables.[[/lang]]
    </li>
    <li>
        [[lang code="en"]]The mean sum of squares is simply the sum of squares per Df.[[/lang]]
        [[lang code="fr"]]Le carré moyen est simplement égal à la somme des carrés divisée par le nombre de degrés de liberté.[[/lang]]
    </li>
    <li>
        [[lang code="en"]]The F value is the ratio of the mean sum of squares and the residual sum of squares. It allows us to compare the variability explained by a variable compared to the unexplained (residual) variability.[[/lang]]
        [[lang code="fr"]]La valeur de F est le rapport entre le carré moyen associé à une variable et le carré moyen résiduel. Elle permet de comparer la variabilité expliquée par une variable à la variabilité non expliquée, c’est-à-dire résiduelle.[[/lang]]
    </li>
</ul>

<p>
</p>

<table>
    <caption></caption>
    <thead>
        <tr>
            <th scope="col" style="text-align: center;border: 1px solid black; " width="100">[[lang code="en"]]Term[[/lang]][[lang code="fr"]]Terme[[/lang]]</th>
            <th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]Df[[/lang]][[lang code="fr"]]ddl[[/lang]]</th>
            <th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]Sum Sq[[/lang]][[lang code="fr"]]Somme des carrés[[/lang]]</th>
            <th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Mean Sq[[/lang]][[lang code="fr"]]Carré moyen[[/lang]]</th>
            <th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]F Value[[/lang]][[lang code="fr"]]Valeur de F[[/lang]]</th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Variety[[/lang]][[lang code="fr"]]Variété[[/lang]]</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@variety_df@}</td>
            <td style="text-align: center;border: 1px solid black; ">{@variety_ss@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@variety_ms@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@variety_f@}</td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Weeding[[/lang]][[lang code="fr"]]Désherbage[[/lang]]</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@weeding_df@}</td>
            <td style="text-align: center;border: 1px solid black; ">{@weeding_ss@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@weeding_ms@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@weeding_f@}</td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Fertiliser[[/lang]][[lang code="fr"]]Engrais[[/lang]]</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@fert_df@}</td>
            <td style="text-align: center;border: 1px solid black; ">{@fert_ss@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@fert_ms@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@fert_f@}</td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Residuals[[/lang]][[lang code="fr"]]Résidus[[/lang]]</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@resid_df@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@resid_ss@}</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@resid_ms@}</td>
            <td style="text-align: center;border: 1px solid black; "></td>
        </tr>
        <tr>
            <td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Total[[/lang]][[lang code="fr"]]Total[[/lang]]</td>
            <td style="color:#00F; text-align: center;border: 1px solid black; ">{@total_df@}</td>
            <td style="text-align: center;border: 1px solid black; ">{@total_ss@}</td>
            <td style="text-align: center;border: 1px solid black; "></td>
            <td style="text-align: center;border: 1px solid black; "></td>
        </tr>
    </tbody>
</table>

<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>dp(x, n) := float(round((x+10^(-8))*10^n)/10^n);

plots: 30+rand(10);
variety: 4+rand(3);
weeding: 2+rand(2);
fert: 3+rand(2);
variety_df:variety-1;
weeding_df: weeding-1;
fert_df: fert-1;
total_df: plots-1;
resid_df: total_df-weeding_df-variety_df-fert_df;

variety_ss: dp(float((rand(10)/10)+(rand(10)+20)),1);
weeding_ss: dp(float((rand(10)/10)+(rand(10)+10)),1);
fert_ss: dp(float((rand(10)/10)+rand(6)+40),1);
resid_ss: dp(float((rand(10)/10)+(rand(20)+65)),1);
total_ss: dp(variety_ss+weeding_ss+fert_ss+resid_ss,1);

variety_ms: dp(variety_ss/variety_df,1);
weeding_ms: dp(weeding_ss/weeding_df,1);
fert_ms: dp(fert_ss/fert_df,1);
resid_ms: dp(resid_ss/resid_df,1);

variety_f: dp(variety_ms/resid_ms,1);
weeding_f: dp(weeding_ms/resid_ms,1);
fert_f: dp(fert_ms/resid_ms,1);

ta1:fert_df;
ta2:resid_df;
ta3:resid_ss;
ta4:variety_ms;
ta5:weeding_ms;
ta6:fert_ms;
ta7:resid_ms;
ta8:variety_f;
ta9:weeding_f;
ta10:fert_f;
ta11: variety_df;
ta12: weeding_df;
ta13: total_df;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans_fert_df</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_fert_f</name>
      <type>numerical</type>
      <tans>ta10</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_fert_ms</name>
      <type>numerical</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_resid_df</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_resid_ms</name>
      <type>numerical</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_resid_ss</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_total_df</name>
      <type>numerical</type>
      <tans>ta13</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans_variety_df</name>
      <type>numerical</type>
      <tans>ta11</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans_variety_f</name>
      <type>numerical</type>
      <tans>ta8</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_variety_ms</name>
      <type>numerical</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_weeding_df</name>
      <type>numerical</type>
      <tans>ta12</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans_weeding_f</name>
      <type>numerical</type>
      <tans>ta9</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans_weeding_ms</name>
      <type>numerical</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <prt>
      <name>prt_fert_df</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fert_df</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]There are {@fert@} levels of fertiliser, hence there are {@ta1@}, i.e. \((n-1)\) degrees of freedom.[[/lang]]
[[lang code="fr"]]Il y a {@fert@} niveaux de fertilisation, donc il y a {@ta1@}, c’est-à-dire \((n-1)\) degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_fert_f</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta10: dp(ans_fert_ms/ans_resid_ms,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans_fert_f</sans>
        <tans>s_ta10</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct F value based on your mean square and residual mean square.[[/lang]]
[[lang code="fr"]]Valeur de F correcte calculée à partir de votre carré moyen et du carré moyen résiduel.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fert_f</sans>
        <tans>ta10</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your F value is correct, but your mean square or residual mean square is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre valeur de F est correcte, mais votre carré moyen ou votre carré moyen résiduel ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]The F value is the ratio of the mean square to the residual mean square.[[/lang]]
[[lang code="fr"]]La valeur de F est le rapport entre le carré moyen et le carré moyen résiduel.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_fert_ms</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta6: dp(fert_ss/ans_fert_df,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fert_ms</sans>
        <tans>s_ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct mean square value based on your sum of squares and Df.[[/lang]]
[[lang code="fr"]]Carré moyen correct calculé à partir de votre somme des carrés et de vos degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
{@fert_ss@} {@ans_fert_df@}<br>
</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fert_ms</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your mean square is correct, but your Df is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre carré moyen est correct, mais votre nombre de degrés de liberté ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Mean square is the sum of squares divided by the degrees of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_resid_df</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta2: ans_total_df-ans_variety_df-ans_weeding_df-ans_fert_df</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_resid_df</sans>
        <tans>s_ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct residual Df based on the total and factor Df values in your table.[[/lang]]
[[lang code="fr"]]Nombre correct de degrés de liberté résiduels calculé à partir des degrés de liberté totaux et de ceux des facteurs dans votre tableau.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_resid_df</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty>ta2</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your residual Df is correct, but some of your other Df values are not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre nombre de degrés de liberté résiduels est correct, mais certains de vos autres degrés de liberté ne le sont pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]The residual is the "left over" after accounting for the factors. Therefore, the residual degrees of freedom are equal to the total degrees of freedom minus the degrees of freedom for the factors.[[/lang]]
[[lang code="fr"]]Le résidu correspond à ce qui reste après avoir pris en compte les facteurs. Ainsi, les degrés de liberté résiduels sont égaux au nombre total de degrés de liberté moins les degrés de liberté des facteurs.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_resid_ms</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta7: dp(ans_resid_ss/ans_resid_df,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans_resid_ms</sans>
        <tans>s_ta7</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct residual mean square value based on your residual sum of squares and residual degrees of freedom.[[/lang]]
[[lang code="fr"]]Carré moyen résiduel correct calculé à partir de votre somme des carrés résiduelle et de vos degrés de liberté résiduels.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans_resid_ms</sans>
        <tans>ta7</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your mean square is correct, but your sum of squares or Df is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre carré moyen est correct, mais votre somme des carrés ou votre nombre de degrés de liberté ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Mean square is the sum of squares divided by the degrees of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_resid_ss</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_resid_ss</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Like the degrees of freedom, the residual sum of squares is the total sum of squares minus the sums of squares for the factors.[[/lang]]
[[lang code="fr"]]Comme pour les degrés de liberté, la somme des carrés résiduelle est égale à la somme des carrés totale moins les sommes des carrés des facteurs.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_variety_f</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta8: dp(ans_variety_ms/ans_resid_ms,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans_variety_f</sans>
        <tans>s_ta8</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct F value based on your mean sum of squares and residual mean square.[[/lang]]
[[lang code="fr"]]Valeur de F correcte calculée à partir de votre carré moyen et du carré moyen résiduel.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans_variety_f</sans>
        <tans>ta8</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your F value is correct, but your mean square is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre valeur de F est correcte, mais votre carré moyen ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]The F value is the ratio of the mean square to the residual mean square.[[/lang]]
[[lang code="fr"]]La valeur de F est le rapport entre le carré moyen et le carré moyen résiduel.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_variety_ms</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta4: dp(variety_ss/ans_variety_df,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_variety_ms</sans>
        <tans>s_ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct mean square value based on your sum of squares and Df.[[/lang]]
[[lang code="fr"]]Carré moyen correct calculé à partir de votre somme des carrés et de vos degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_variety_ms</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your mean sum of squares is correct, but your Df is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre carré moyen est correct, mais votre nombre de degrés de liberté ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Mean sum of squares is the sum of squares per degree of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_weeding_df</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_weeding_df</sans>
        <tans>ta12</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]There are {@weeding@} levels of weeding, hence there are {@ta12@}, i.e. \((n-1)\) degrees of freedom.[[/lang]]
[[lang code="fr"]]Il y a {@weeding@} niveaux de désherbage, donc il y a {@ta12@}, c’est-à-dire \((n-1)\) degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_weeding_f</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta9: dp(ans_weeding_ms/ans_resid_ms,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans_weeding_f</sans>
        <tans>s_ta9</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct F value based on your mean square and residual mean square.[[/lang]]
[[lang code="fr"]]Valeur de F correcte calculée à partir de votre carré moyen et du carré moyen résiduel.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans_weeding_f</sans>
        <tans>ta9</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your F value is correct, but your mean square or residual mean square is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre valeur de F est correcte, mais votre carré moyen ou votre carré moyen résiduel ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]The F value is the ratio of the mean square to the residual mean square.[[/lang]]
[[lang code="fr"]]La valeur de F est le rapport entre le carré moyen et le carré moyen résiduel.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_weeding_ms</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta5: dp(weeding_ss/ans_weeding_df,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_weeding_ms</sans>
        <tans>s_ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Correct mean square value based on your sum of squares and Df.[[/lang]]
[[lang code="fr"]]Carré moyen correct calculé à partir de votre somme des carrés et de vos degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_weeding_ms</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Hmm, how did this happen? Your mean square is correct, but your Df is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre carré moyen est correct, mais votre nombre de degrés de liberté ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]
<br>
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Mean square is the sum of squares divided by the degrees of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1_total_df</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_total_df</sans>
        <tans>ta13</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt13-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt13-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]The total degrees of freedom is the number of observations minus \(1\), so {@total_df@}.[[/lang]]
[[lang code="fr"]]Le nombre total de degrés de liberté est égal au nombre d’observations moins \(1\), soit {@total_df@}.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1_variety_df</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_variety_df</sans>
        <tans>ta11</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]There are {@variety@} different varieties, hence there are {@ta11@}, i.e. \((n-1)\) degrees of freedom.[[/lang]]
[[lang code="fr"]]Il y a {@variety@} variétés différentes, donc il y a {@ta11@}, c’est-à-dire \((n-1)\) degrés de liberté.[[/lang]]
<br>
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
