<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>9: Interpreting ANOVA</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>[[lang code="en"]]An experiment is conducted on {@plots@} plots to understand the effect of several factors on the yield of millet.[[/lang]]
[[lang code="fr"]]Une expérience est réalisée sur {@plots@} parcelles afin d'étudier l'effet de plusieurs facteurs sur le rendement du millet.[[/lang]]</p>

<p>[[lang code="en"]]The factors are: {@variety@} varieties of millet, {@weeding@} types of weeding and {@fert@} different levels of fertiliser.[[/lang]]
[[lang code="fr"]]Les facteurs étudiés sont : {@variety@} variétés de millet, {@weeding@} types de désherbage et {@fert@} niveaux de fertilisation.[[/lang]]</p>

<p><b>
[[lang code="en"]]Give the values to \(1\) decimal place.[[/lang]]
[[lang code="fr"]]Donnez les valeurs à \(1\) décimale près.[[/lang]]
</b></p>

<br>

<table>
    <thead>
        <tr>
            <th scope="col" width="100">[[lang code="en"]]Term[[/lang]][[lang code="fr"]]Terme[[/lang]]</th>
            <th scope="col" width="100">[[lang code="en"]]Df[[/lang]][[lang code="fr"]]ddl[[/lang]]</th>
            <th scope="col" width="100">[[lang code="en"]]Sum Sq[[/lang]][[lang code="fr"]]Somme des carrés[[/lang]]</th>
            <th scope="col" width="100">[[lang code="en"]]Mean Sq[[/lang]][[lang code="fr"]]Carré moyen[[/lang]]</th>
            <th scope="col" width="100">[[lang code="en"]]F Value[[/lang]][[lang code="fr"]]Valeur de F[[/lang]]</th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td>[[lang code="en"]]Variety[[/lang]][[lang code="fr"]]Variété[[/lang]]</td>
            <td>{@variety_df@}<br></td>
            <td>{@variety_ss@}</td>
            <td>{@variety_ms@}</td>
            <td>{@variety_f@}</td>
        </tr>
        <tr>
            <td>[[lang code="en"]]Weeding[[/lang]][[lang code="fr"]]Désherbage[[/lang]]</td>
            <td>{@weeding_df@}<br></td>
            <td>{@weeding_ss@}</td>
            <td>{@weeding_ms@}</td>
            <td>{@weeding_f@}</td>
        </tr>
        <tr>
            <td>[[lang code="en"]]Fertiliser[[/lang]][[lang code="fr"]]Fertilisation[[/lang]]</td>
            <td>{@fert_df@}</td>
            <td>{@fert_ss@}</td>
            <td>{@fert_ms@}<br></td>
            <td>{@fert_f@}</td>
        </tr>
        <tr>
            <td>[[lang code="en"]]Residuals[[/lang]][[lang code="fr"]]Résidus[[/lang]]</td>
            <td>{@resid_df@}</td>
            <td>{@resid_ss@}</td>
            <td>{@resid_ms@}</td>
            <td></td>
        </tr>
        <tr>
            <td>[[lang code="en"]]Total[[/lang]][[lang code="fr"]]Total[[/lang]]</td>
            <td>{@total_df@}<br></td>
            <td>{@total_ss@}</td>
            <td></td>
            <td></td>
        </tr>
    </tbody>
</table>

<br>

<p>
[[lang code="en"]]What percentage (to one decimal place) of the variation in yield is explained by differences in variety?[[/lang]]
[[lang code="fr"]]Quel pourcentage (à une décimale près) de la variation du rendement est expliqué par les différences entre les variétés ?[[/lang]]
</p>

<p>[[input:ans_variety]] [[validation:ans_variety]] [[feedback:prt1_variety]]</p>

<p>
[[lang code="en"]]What percentage (to one decimal place) of the variation in yield is not accounted for?[[/lang]]
[[lang code="fr"]]Quel pourcentage (à une décimale près) de la variation du rendement n'est pas expliqué par le modèle ?[[/lang]]
</p>

<p>[[input:ans_unaccounted]] [[validation:ans_unaccounted]] [[feedback:prt1_unaccounted]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>dp(x, n) := float(round((x+10^(-8))*10^n)/10^n);

plots: 30+rand(10);
variety: 4+rand(3);
weeding: 2+rand(2);
fert: 3+rand(2);
variety_df:variety-1;
weeding_df: weeding-1;
fert_df: fert-1;
total_df: plots-1;
resid_df: total_df-weeding_df-variety_df-fert_df;

variety_ss: dp(float((rand(10)/10)+(rand(10)+20)),1);
weeding_ss: dp(float((rand(10)/10)+(rand(10)+10)),1);
fert_ss: dp(float((rand(10)/10)+rand(6)+40),1);
resid_ss: dp(float((rand(10)/10)+(rand(20)+65)),1);
total_ss: dp(variety_ss+weeding_ss+fert_ss+resid_ss,1);

variety_ms: dp(variety_ss/variety_df,1);
weeding_ms: dp(weeding_ss/weeding_df,1);
fert_ms: dp(fert_ss/fert_df,1);
resid_ms: dp(resid_ss/resid_df,1);

variety_f: dp(variety_ms/resid_ms,1);
weeding_f: dp(weeding_ms/resid_ms,1);
fert_f: dp(fert_ms/resid_ms,1);

variety_variation: dp(variety_ss/total_ss * 100, 1)
unaccounted_variation: dp(resid_ss/total_ss * 100, 1)

ta1: variety_variation;
ta2: unaccounted_variation;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans_unaccounted</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans_variety</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1_unaccounted</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_unaccounted</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1_unaccounted-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1_unaccounted-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]The unexplained variation can be found by dividing the residual sum of squares by the total sum of squares, then multiplying by 100 to obtain the percentage.[[/lang]]
[[lang code="fr"]]La part de variation non expliquée s'obtient en divisant la somme des carrés résiduelle par la somme des carrés totale, puis en multipliant le résultat par 100 pour obtenir un pourcentage.[[/lang]]
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1_variety</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_variety</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1_variety-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">
[[lang code="en"]]Variety accounts for {@variety_ss@}/{@total_ss@} × 100% of the variation in yield.[[/lang]]
[[lang code="fr"]]La variété explique {@variety_ss@}/{@total_ss@} × 100&nbsp;% de la variation du rendement.[[/lang]]
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
