<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Fill ANOVA table for 2-year tomato data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
[[lang code="en"]]An experiment was conducted over two years to investigate the effects of three factors—light, heat, and variety—each at two levels, on tomato yields.[[/lang]]
[[lang code="fr"]]Une expérience a été menée pendant deux ans afin d'étudier les effets de trois facteurs — la lumière, la température et la variété — chacun à deux niveaux, sur le rendement des tomates.[[/lang]]
</p>

<p>
[[lang code="en"]]Here is an incomplete analysis of variance (ANOVA) table for the tomato data for the <b>\(2\) years together</b>.[[/lang]]
[[lang code="fr"]]Voici un tableau d'analyse de la variance (ANOVA) incomplet pour les données sur les tomates correspondant aux <b>\(2\) années combinées</b>.[[/lang]]
</p>

<p>
[[lang code="en"]]Complete the ANOVA table.[[/lang]]
[[lang code="fr"]]Complétez le tableau d'ANOVA.[[/lang]]
</p>

<p><b>
[[lang code="en"]]Give the values to \(1\) decimal place.[[/lang]]
[[lang code="fr"]]Donnez les valeurs à \(1\) décimale près.[[/lang]]
</b></p>

<p>
<style type="text/css">
#maintable td.blue {color:#00F;}
</style>

<table>
<caption></caption>
<thead>
<tr>
<th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Term[[/lang]][[lang code="fr"]]Terme[[/lang]]</th>
<th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]Df[[/lang]][[lang code="fr"]]ddl[[/lang]]</th>
<th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]Sum Sq[[/lang]][[lang code="fr"]]Somme des carrés[[/lang]]</th>
<th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Mean Sq[[/lang]][[lang code="fr"]]Carré moyen[[/lang]]</th>
<th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]F Value[[/lang]][[lang code="fr"]]Valeur de F[[/lang]]</th>
</tr>
</thead>
<tbody>
<tr>
<td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Light[[/lang]][[lang code="fr"]]Lumière[[/lang]]</td>
<td style="text-align: center;border: 1px solid black;">{@light_df@}</td>
<td style="text-align: center;border: 1px solid black;">{@light_ss@}</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans8]] [[validation:ans8]] [[feedback:prt8]]</td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Heat[[/lang]][[lang code="fr"]]Température[[/lang]]</td>
<td style="text-align: center;border: 1px solid black;">{@heat_df@}</td>
<td style="text-align: center;border: 1px solid black;">{@heat_ss@}</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans9]] [[validation:ans9]] [[feedback:prt9]]</td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Variety[[/lang]][[lang code="fr"]]Variété[[/lang]]</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans1]] [[validation:ans1]] [[feedback:prt1a]]</td>
<td style="text-align: center;border: 1px solid black;">{@variety_ss@}</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans10]] [[validation:ans10]] [[feedback:prt10]]</td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Residuals[[/lang]][[lang code="fr"]]Résidus[[/lang]]</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</td>
<td style="text-align: center;border: 1px solid black;">[[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</td>
<td style="text-align: center;border: 1px solid black;"></td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black;">[[lang code="en"]]Total[[/lang]][[lang code="fr"]]Total[[/lang]]</td>
<td style="text-align: center;border: 1px solid black;">{@total_df@}</td>
<td style="text-align: center;border: 1px solid black;">{@total_ss@}</td>
<td style="text-align: center;border: 1px solid black;"></td>
<td style="text-align: center;border: 1px solid black;"></td>
</tr>
</tbody>
</table>
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>
[[lang code="en"]]The complete ANOVA table is shown below.[[/lang]]
[[lang code="fr"]]Le tableau d’ANOVA complet est présenté ci-dessous.[[/lang]]
</p>

<p>
[[lang code="en"]]There are a few key points to remember:[[/lang]]
[[lang code="fr"]]Quelques points essentiels sont à retenir :[[/lang]]
</p>

<ul>
  <li>
    [[lang code="en"]]The residual is the "left over" after taking into account the variables. So the residual Df and sum of squares are the difference between the total values and the sum of the values for the variables.[[/lang]]
    [[lang code="fr"]]Le résidu correspond à ce qui reste après avoir pris en compte les variables. Ainsi, les degrés de liberté et la somme des carrés résiduels correspondent à la différence entre les valeurs totales et la somme des valeurs associées aux variables.[[/lang]]
  </li>
  <li>
    [[lang code="en"]]The mean square is simply the sum of squares divided by the degrees of freedom.[[/lang]]
    [[lang code="fr"]]Le carré moyen est simplement égal à la somme des carrés divisée par le nombre de degrés de liberté.[[/lang]]
  </li>
  <li>
    [[lang code="en"]]The F value is the ratio of the mean square to the residual mean square. It allows us to compare the variability explained by a variable with the unexplained (residual) variability.[[/lang]]
    [[lang code="fr"]]La valeur de F est le rapport entre le carré moyen et le carré moyen résiduel. Elle permet de comparer la variabilité expliquée par une variable à la variabilité non expliquée, c’est-à-dire résiduelle.[[/lang]]
  </li>
</ul>

<p>
<table>
<caption></caption>
<thead>
<tr>
<th scope="col" style="text-align: center;border: 1px solid black; " width="100">[[lang code="en"]]Term[[/lang]][[lang code="fr"]]Terme[[/lang]]</th>
<th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]Df[[/lang]][[lang code="fr"]]ddl[[/lang]]</th>
<th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]Sum Sq[[/lang]][[lang code="fr"]]Somme des carrés[[/lang]]</th>
<th scope="col" style="text-align: center;border: 1px solid black;" width="100">[[lang code="en"]]Mean Sq[[/lang]][[lang code="fr"]]Carré moyen[[/lang]]</th>
<th scope="col" style="text-align: center; border: 1px solid black;" width="100">[[lang code="en"]]F Value[[/lang]][[lang code="fr"]]Valeur de F[[/lang]]</th>
</tr>
</thead>
<tbody>
<tr>
<td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Light[[/lang]][[lang code="fr"]]Lumière[[/lang]]</td>
<td class="blue" style="text-align: center;border: 1px solid black; ">{@light_df@}</td>
<td style="text-align: center;border: 1px solid black; ">{@light_ss@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta4@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta8@}</td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Heat[[/lang]][[lang code="fr"]]Température[[/lang]]</td>
<td style="text-align: center;border: 1px solid black; ">{@heat_df@}</td>
<td style="text-align: center;border: 1px solid black; ">{@heat_ss@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta5@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta9@}</td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Variety[[/lang]][[lang code="fr"]]Variété[[/lang]]</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta1@}</td>
<td style="text-align: center;border: 1px solid black; ">{@variety_ss@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta6@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta10@}</td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Residuals[[/lang]][[lang code="fr"]]Résidus[[/lang]]</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta2@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta3@}</td>
<td style="color:#00F; text-align: center;border: 1px solid black; ">{@ta7@}</td>
<td style="text-align: center;border: 1px solid black; "></td>
</tr>
<tr>
<td style="text-align: center;border: 1px solid black; ">[[lang code="en"]]Total[[/lang]][[lang code="fr"]]Total[[/lang]]</td>
<td style="text-align: center;border: 1px solid black; ">{@total_df@}</td>
<td style="text-align: center;border: 1px solid black; ">{@total_ss@}</td>
<td style="text-align: center;border: 1px solid black; "></td>
<td style="text-align: center;border: 1px solid black; "></td>
</tr>
</tbody>
</table>
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>dp(x, n) := float(round((x+10^(-8))*10^n)/10^n);

light_df:1 ;
heat_df: 1;
variety_df: 1;
total_df: 23;
resid_df: total_df-heat_df-light_df-variety_df;

light_ss: dp(float((rand(10)/10)+(rand(6)+5)),1);
heat_ss: dp(float((rand(10)/10)+(rand(10)+40)),1);
variety_ss: dp(float((rand(10)/10)+(rand(6)+3)+light_ss),1);
resid_ss: dp(float((rand(10)/10)+(rand(20)+95)),1);
total_ss: dp(light_ss+heat_ss+variety_ss+resid_ss,1);

light_ms: dp(light_ss/light_df,1);
heat_ms: dp(heat_ss/heat_df,1);
variety_ms: dp(variety_ss/variety_df,1);
resid_ms: dp(resid_ss/resid_df,1);

light_f: dp(light_ms/resid_ms,1);
heat_f: dp(heat_ms/resid_ms,1);
variety_f: dp(variety_ms/resid_ms,1);

ta1:variety_df;
ta2:resid_df;
ta3:resid_ss;
ta4:light_ms;
ta5:heat_ms;
ta6:variety_ms;
ta7:resid_ms;
ta8:light_f;
ta9:heat_f;
ta10:variety_f;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans10</name>
      <type>numerical</type>
      <tans>ta10</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans5</name>
      <type>numerical</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans6</name>
      <type>numerical</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans7</name>
      <type>numerical</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans8</name>
      <type>numerical</type>
      <tans>ta8</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <input>
      <name>ans9</name>
      <type>numerical</type>
      <tans>ta9</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>maxdp:1</options>
    </input>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta10: dp(ans6/ans7,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>s_ta10</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Correct F value based on your mean square and residual mean square.[[/lang]]
[[lang code="fr"]]Valeur de F correcte calculée à partir de votre carré moyen et du carré moyen résiduel.[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>ta10</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Hmm, how did this happen? Your F value is correct, but your mean square or residual mean square is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre valeur de F est correcte, mais votre carré moyen ou votre carré moyen résiduel ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]The F value is the ratio of the mean square to the residual mean square.[[/lang]]
[[lang code="fr"]]La valeur de F est le rapport entre le carré moyen et le carré moyen résiduel.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1a</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]In the tomato data there are two varieties, hence there is \(1\), i.e. \((n-1)\), degree of freedom.[[/lang]]
[[lang code="fr"]]Dans les données sur les tomates, il y a deux variétés ; il y a donc \(1\), c’est-à-dire \((n-1)\), degré de liberté.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta2: total_df-light_df-heat_df-ans1</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>s_ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Correct residual degrees of freedom based on the degrees of freedom for light, heat and variety in your table.[[/lang]]
[[lang code="fr"]]Nombre correct de degrés de liberté résiduels calculé à partir des degrés de liberté de la lumière, de la température et de la variété dans votre tableau.[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Hmm, how did this happen? Your residual degrees of freedom are correct, but your variety degrees of freedom are not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre nombre de degrés de liberté résiduels est correct, mais celui de la variété ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]The residual is the "left over" after accounting for the factors. Therefore, the residual degrees of freedom are equal to the total degrees of freedom minus the degrees of freedom for the factors.[[/lang]]
[[lang code="fr"]]Le résidu correspond à ce qui reste après avoir pris en compte les facteurs. Ainsi, les degrés de liberté résiduels sont égaux au nombre total de degrés de liberté moins les degrés de liberté des facteurs.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Like the degrees of freedom, the residual sum of squares is the total sum of squares minus the sums of squares for the factors.[[/lang]]
[[lang code="fr"]]Comme pour les degrés de liberté, la somme des carrés résiduelle est égale à la somme des carrés totale moins les sommes des carrés des facteurs.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Mean square is the sum of squares divided by the degrees of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Mean square is the sum of squares divided by the degrees of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta6: dp(variety_ss/ans1,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>s_ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Correct mean square value based on your sum of squares and degrees of freedom.[[/lang]]
[[lang code="fr"]]Carré moyen correct calculé à partir de votre somme des carrés et de vos degrés de liberté.[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Hmm, how did this happen? Your mean square is correct, but your Df is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre carré moyen est correct, mais votre nombre de degrés de liberté ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Mean square is the sum of squares divided by the degrees of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta7: dp(ans3/ans2,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>s_ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Correct residual mean square based on your residual sum of squares and residual degrees of freedom.[[/lang]]
[[lang code="fr"]]Carré moyen résiduel correct calculé à partir de votre somme des carrés résiduelle et de vos degrés de liberté résiduels.[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Hmm, how did this happen? Your mean square is correct, but your sum of squares or Df is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre carré moyen est correct, mais votre somme des carrés ou votre nombre de degrés de liberté ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Mean square is the sum of squares divided by the degrees of freedom.[[/lang]]
[[lang code="fr"]]Le carré moyen est obtenu en divisant la somme des carrés par le nombre de degrés de liberté.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta8: dp(ans4/ans7,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>s_ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Correct F value based on your mean square and residual mean square.[[/lang]]
[[lang code="fr"]]Valeur de F correcte calculée à partir de votre carré moyen et du carré moyen résiduel.[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Hmm, how did this happen? Your F value is correct, but your mean square is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre valeur de F est correcte, mais votre carré moyen ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]The F value is the ratio of the mean square to the residual mean square.[[/lang]]
[[lang code="fr"]]La valeur de F est le rapport entre le carré moyen et le carré moyen résiduel.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ta9: dp(ans5/ans7,1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>s_ta9</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Correct F value based on your mean square and residual mean square.[[/lang]]
[[lang code="fr"]]Valeur de F correcte calculée à partir de votre carré moyen et du carré moyen résiduel.[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>ta9</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[[[lang code="en"]]Hmm, how did this happen? Your F value is correct, but your mean square or residual mean square is not. Did you make a typing mistake? Anyway, you still get some marks :)[[/lang]]
[[lang code="fr"]]Hum, comment cela a-t-il pu arriver ? Votre valeur de F est correcte, mais votre carré moyen ou votre carré moyen résiduel ne l’est pas. Avez-vous fait une erreur de saisie ? Quoi qu’il en soit, vous obtenez tout de même quelques points :)[[/lang]]]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[[[lang code="en"]]The F value is the ratio of the mean square to the residual mean square.[[/lang]]
[[lang code="fr"]]La valeur de F est le rapport entre le carré moyen et le carré moyen résiduel.[[/lang]]]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
