<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>5: Appending Data</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>
[[lang code='en,other']]
Rainfall data were taken in two regions in Ghana: Elamat and Palstond. However, the data have been collected into two data frames: one contains data from 1990–1999, and the other from 2000–2016.
[[/lang]]
[[lang code='fr']]
Des données de précipitations ont été relevées dans deux régions du Ghana : Elamat et Palstond. Cependant, les données ont été regroupées dans deux tableaux de données : l’un contient les données de 1990 à 1999, et l’autre celles de 2000 à 2016.
[[/lang]]
</p>

<p>
[[lang code='en,other']]
The former data set contains the station name, date, year, rainfall (mm), minimum temperature (Celsius), and maximum temperature (Celsius).
[[/lang]]
[[lang code='fr']]
Le premier jeu de données contient le nom de la station, la date, l’année, les précipitations (mm), la température minimale (en degrés Celsius) et la température maximale (en degrés Celsius).
[[/lang]]
</p>

<p>
[[lang code='en,other']]
The latter data set contains the station name, date, year, rainfall (cm), minimum temperature (Celsius), maximum temperature (Celsius), and hours of sunshine.
[[/lang]]
[[lang code='fr']]
Le second jeu de données contient le nom de la station, la date, l’année, les précipitations (cm), la température minimale (en degrés Celsius), la température maximale (en degrés Celsius) et le nombre d’heures d’ensoleillement.
[[/lang]]
</p>

<p>
[[lang code='en,other']]
The task is to combine the data sets and answer the following questions:
<br><em>(Note that these are made-up data.)</em>
[[/lang]]
[[lang code='fr']]
L’objectif est de combiner les deux jeux de données et de répondre aux questions suivantes :
<br><em>(Remarque : ces données sont fictives.)</em>
[[/lang]]
</p>

<p>
  <button type="button"
          class="btn btn-primary"
          onclick="downloadDataset(
            'https://raw.githubusercontent.com/IDEMSInternational/statistics-textbooks-questions/refs/heads/master/question-datasets/ghana-1990-1999.csv',
            'ghana-1990-1999.csv'
          )">
    [[lang code='en,other']]Download Data: 1990–1999[[/lang]]
    [[lang code='fr']]Télécharger les données : 1990–1999[[/lang]]
  </button>
</p>

<p>
  <button type="button"
          class="btn btn-primary"
          onclick="downloadDataset(
            'https://raw.githubusercontent.com/IDEMSInternational/statistics-textbooks-questions/refs/heads/master/question-datasets/ghana-2000-2016.csv',
            'ghana-2000-2016.csv'
          )">
    [[lang code='en,other']]Download Data: 2000–2016[[/lang]]
    [[lang code='fr']]Télécharger les données : 2000–2016[[/lang]]
  </button>
</p>

<br>

<p>
[[lang code='en,other']]
1. Which year had the least overall rainfall?
[[/lang]]
[[lang code='fr']]
1. Quelle année a enregistré le plus faible total de précipitations ?
[[/lang]]
</p>

<p>[[input:ans1]] [[validation:ans1]]</p>

<br>

<p>
[[lang code='en,other']]
2. Which region had more rainfall between 1990 and 2016, inclusive?
[[/lang]]
[[lang code='fr']]
2. Quelle région a enregistré le plus de précipitations entre 1990 et 2016 inclus ?
[[/lang]]
</p>

<p>[[input:ans2]] [[validation:ans2]]</p>

<br>

<p>
[[lang code='en,other']]
3. How much rainfall was there in total across both stations in 2016, in \(mm\)?
<br><em>Give your answer to 2 decimal places.</em>
[[/lang]]
[[lang code='fr']]
3. Quelle quantité totale de précipitations a été enregistrée dans les deux stations en 2016, en \(mm\) ?
<br><em>Donnez votre réponse avec 2 décimales.</em>
[[/lang]]
</p>

<p>[[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt1]]</p>

<script>
async function downloadDataset(url, filename) {
    try {
        const response = await fetch(url);

        if (!response.ok) {
            throw new Error(
                'The dataset could not be retrieved. HTTP status: ' +
                response.status
            );
        }

        const blob = await response.blob();
        const objectUrl = URL.createObjectURL(blob);

        const downloadLink = document.createElement('a');
        downloadLink.href = objectUrl;
        downloadLink.download = filename;
        downloadLink.style.display = 'none';

        document.body.appendChild(downloadLink);
        downloadLink.click();
        downloadLink.remove();

        setTimeout(function () {
            URL.revokeObjectURL(objectUrl);
        }, 1000);
    } catch (error) {
        console.error(error);

        alert(
            'The file could not be downloaded automatically. ' +
            'The dataset will now be opened in a new tab.'
        );

        window.open(url, '_blank', 'noopener,noreferrer');
    }
}
</script>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ma1: [
  [1990, false],  [1991, false],  [1992, true],  [1993, false],
  [1994, false],  [1995, false],  [1996, false],  [1997, false],
  [1998, false],  [1999, false],  [2000, false],  [2001, false],
  [2002, false],  [2003, false],  [2004, false],  [2005, false],
  [2006, false],  [2007, false],  [2008, false],  [2009, false],
  [2010, false],  [2011, false],  [2012, false],  [2013, false],
  [2014, false],  [2015, false],  [2016, false]
];

ma2: [
  ["Palstond", false],
  ["Elamat", true],
  ["Unknown", false]
];

ta3: 1869.80;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ma1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ma2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* node 1: check answers correct */
node1_test: if ans1=1992 and ans2="Elamat" and ans3=1869.80 then 1 else 0;

/* node 2: fails to compare the units in each question */
node2_test: if ans1=2013 and ans2="Palstond" and ans3=186.98 then 1 else 0;

/* node 3: treats 2000-16 stations as unknown */
node3_test: if ans1=1992 and member(ans2, ["Palstond", "Unknown"]) and ans3=1869.80 then 1 else 0;

/* node 4: disregards the order of the columns */
node4_test: if ans1=2013 and ans2="Elamat" and member(ans3, [1869.8, 186.98]) then 1 else 0;

/* node 5: disregards the order of the columns */
node5_test: if ans1=1992 and ans2="Palstond" and member(ans3, [17567.7, 175677]) then 1 else 0;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check correct answers</description>
        <answertest>AlgEquiv</answertest>
        <sans>node1_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
[[lang code='en,other']]
Great! It appears that when combining the data sets, you have taken into account the different units used!
<br>In addition, it seems that you have correctly appended the data sets together.
[[/lang]]
[[lang code='fr']]
Bravo ! Il semble que, lors de la combinaison des jeux de données, vous ayez bien tenu compte des différentes unités utilisées.
<br>De plus, il semble que vous ayez correctement concaténé les deux jeux de données.
[[/lang]]
</p>

<br>

<p>
[[lang code='en,other']]
There is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
[[/lang]]
[[lang code='fr']]
Une solution détaillée de cet exercice est disponible <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">ici</a>.
[[/lang]]
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>fails to compare the units in each question</description>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
[[lang code='en,other']]
It appears that when combining the data sets, you have not taken into account the different units used.
<br>Have a look at a plot of your rainfall by year. Does anything seem odd here?
<br>We mentioned in the initial question that rainfall is measured in cm in one data set and mm in another.
[[/lang]]
[[lang code='fr']]
Il semble que, lors de la combinaison des jeux de données, vous n'ayez pas tenu compte des différentes unités utilisées.
<br>Examinez le graphique des précipitations par année. Remarquez-vous quelque chose d'inhabituel ?
<br>Comme indiqué dans l'énoncé, les précipitations sont mesurées en cm dans un jeu de données et en mm dans l'autre.
[[/lang]]
</p>

<br>

<p style="margin-left:30px">
  <img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-1.jpg?raw=true">
</p>

<br>

<p>
[[lang code='en,other']]
Try again, but if you are stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
[[/lang]]
[[lang code='fr']]
Essayez à nouveau. Si vous êtes bloqué(e), une solution détaillée de cet exercice est disponible <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">ici</a>.
[[/lang]]
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>treats 2000-16 stations as unknown</description>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
[[lang code='en,other']]
It appears that when combining the data sets, you have taken into account the different units used, which is great!
<br>However, have a look at a plot of rainfall by year, including the station. Does anything seem odd here?
<br>All stations from the year 2000 onwards are shown as unknown. Have a look at the station variable name in the two initial data sets. Can you see why this might have happened when combining the data?
[[/lang]]
[[lang code='fr']]
Il semble que, lors de la combinaison des jeux de données, vous ayez bien tenu compte des différentes unités utilisées, ce qui est très bien !
<br>Cependant, examinez un graphique des précipitations par année, en distinguant les stations. Remarquez-vous quelque chose d’inhabituel ?
<br>À partir de l’année 2000, toutes les stations sont indiquées comme étant inconnues. Examinez le nom de la variable correspondant à la station dans les deux jeux de données initiaux. Comprenez-vous pourquoi cela a pu se produire lors de la combinaison des données ?
[[/lang]]
</p>

<br>

<p style="margin-left:30px">
  <img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-2.jpg?raw=true"
       alt="Plot of rainfall by year and station">
</p>

<br>

<p>
[[lang code='en,other']]
Try again, but if you are stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
[[/lang]]
[[lang code='fr']]
Essayez à nouveau. Si vous êtes bloqué(e), une solution détaillée de cet exercice est disponible <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">ici</a>.
[[/lang]]
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>disregards the order of the columns</description>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
[[lang code='en,other']]
It appears that when combining the data sets, you have not taken into account the order of the columns.
<br>Have a look at the data: the columns are arranged in a different order in the two data sets.
<br>This can be seen more clearly in a plot of rainfall by year.
<br>We want to append the two data sheets. In R-Instat, go to <strong>Prepare &gt; Data Reshape &gt; Append</strong>. We suggest reviewing the tutorial on merging and appending data in R-Instat.
[[/lang]]
[[lang code='fr']]
Il semble que, lors de la combinaison des jeux de données, vous n’ayez pas tenu compte de l’ordre des colonnes.
<br>Examinez les données : les colonnes sont disposées dans un ordre différent dans les deux jeux de données.
<br>Cela apparaît plus clairement dans un graphique des précipitations par année.
<br>Nous souhaitons concaténer les deux feuilles de données. Dans R-Instat, allez dans <strong>Prepare &gt; Data Reshape &gt; Append</strong>. Nous vous conseillons de consulter le tutoriel sur la fusion et la concaténation de données dans R-Instat.
[[/lang]]
</p>

<br>

<p style="margin-left:30px">
  <img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-4.jpg?raw=true"
       alt="Plot showing the effect of appending columns in the wrong order">
</p>

<br>

<p>
[[lang code='en,other']]
Try again, but if you are stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
[[/lang]]
[[lang code='fr']]
Essayez à nouveau. Si vous êtes bloqué(e), une solution détaillée de cet exercice est disponible <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">ici</a>.
[[/lang]]
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>disregards the order of the columns</description>
        <answertest>AlgEquiv</answertest>
        <sans>node5_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
[[lang code='en,other']]
It appears that when combining the data sets, you have not taken into account the order of the columns.
<br>Have a look at the data: the columns are arranged in a different order in the two data sets.
<br>This can be seen more clearly in a plot of rainfall by year.
<br>We want to append the two data sheets. In R-Instat, go to <strong>Prepare &gt; Data Reshape &gt; Append</strong>. We suggest reviewing the tutorial on merging and appending data in R-Instat.
[[/lang]]
[[lang code='fr']]
Il semble que, lors de la combinaison des jeux de données, vous n’ayez pas tenu compte de l’ordre des colonnes.
<br>Examinez les données : les colonnes sont disposées dans un ordre différent dans les deux jeux de données.
<br>Cela apparaît plus clairement sur un graphique des précipitations par année.
<br>Nous souhaitons concaténer les deux feuilles de données. Dans R-Instat, allez dans <strong>Prepare &gt; Data Reshape &gt; Append</strong>. Nous vous conseillons de consulter le tutoriel sur la fusion et la concaténation de données dans R-Instat.
[[/lang]]
</p>

<br>

<p style="margin-left:30px">
  <img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-3.jpg?raw=true"
       alt="Plot showing the effect of appending columns in the wrong order">
</p>

<p>
[[lang code='en,other']]
Try again, but if you are stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
[[/lang]]
[[lang code='fr']]
Essayez à nouveau. Si vous êtes bloqué(e), une solution détaillée de cet exercice est disponible <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">ici</a>.
[[/lang]]
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>
[[lang code='en,other']]
Not quite&mdash;are you looking at the <b>overall rainfall</b> for each year?
<br>We want to append the two data sheets. In R-Instat, go to <strong>Prepare &gt; Data Reshape &gt; Append</strong>. We suggest reviewing the tutorial on merging and appending data in R-Instat.
[[/lang]]
[[lang code='fr']]
Pas tout à fait. Avez-vous bien pris en compte les <b>précipitations totales</b> pour chaque année ?
<br>Nous souhaitons concaténer les deux feuilles de données. Dans R-Instat, allez dans <strong>Prepare &gt; Data Reshape &gt; Append</strong>. Nous vous conseillons de consulter le tutoriel sur la fusion et la concaténation de données dans R-Instat.
[[/lang]]
</p>

<br>

<p>
[[lang code='en,other']]
Try again, but if you are stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
[[/lang]]
[[lang code='fr']]
Essayez à nouveau. Si vous êtes bloqué(e), une solution détaillée de cet exercice est disponible <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">ici</a>.
[[/lang]]
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
