<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>5: Appending Data</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Rainfall data were taken in two regions in Ghana: Elamat and Palstond. However, the data has been collected into two data frames: one contains data from 1990-1999, and the other from 2000-2016.</p>

<p>The former data set contains the station name, the date, year, rainfall (mm), minimum temperature (Celsius), and maximum temperature (Celsius).</p>

<p>The latter data set contains the station name, date, year, rainfall (cm), minimum temperature (Celsius), and maximum temperature (Celsius), and hours of sunshine.</p>

<p>The task is to combine the data sets and answer the following questions:
<br><em>(Note that this is “made up” data)</em></p>

<p><button onclick="exportToCsv1990()">Download Data: 1990-1999</button></p>
<p><button onclick="exportToCsv2000()">Download Data: 2000-2016</button></p>

<br>

<p>1. Which year had the least overall rainfall?</p>
<p>[[input:ans1]] [[validation:ans1]]</p>

<br>

<p>2. Which region had more rainfall between 1990-2016 (inclusive)?</p>
<p>[[input:ans2]] [[validation:ans2]]</p>

<br>

<p>3. How much rainfall was there in total across both stations in 2016 in \(mm\)?</p>
<br><em>Give your answer to 2 decimal places.</em></p>
<p>[[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt1]]</p>

<script>
var Results1990;
var Results2000;

fetch("https://raw.githubusercontent.com/IDEMSInternational/wxmaxima-stack/main/data-files/ghana-1990-1999.txt")
  .then(res => res.json())
  .then(data => {
    Results1990 = data;
   })
fetch("https://raw.githubusercontent.com/IDEMSInternational/wxmaxima-stack/main/data-files/ghana-2000-2016.txt")
  .then(res => res.json())
  .then(data => {
    Results2000 = data;
   })

exportToCsv1990 = function(file_name) {
  var CsvString = "";
  Results1990.forEach(function(RowItem, RowIndex) {
    RowItem.forEach(function(ColItem, ColIndex) {
             CsvString += ColItem;
             if( ColIndex < RowItem.length - 1){
                 CsvString += ',';
                } 
    });
    CsvString += "\r\n";
  });
  CsvString = "data:application/csv," + encodeURIComponent(CsvString);
 var x = document.createElement("A");
 x.setAttribute("href", CsvString );
 x.setAttribute("download","data_set.csv");
 document.body.appendChild(x);
 x.click();
}

exportToCsv2000 = function(file_name) {
  var CsvString = "";
  Results2000.forEach(function(RowItem, RowIndex) {
    RowItem.forEach(function(ColItem, ColIndex) {
             CsvString += ColItem;
             if( ColIndex < RowItem.length - 1){
                 CsvString += ',';
                } 
    });
    CsvString += "\r\n";
  });
  CsvString = "data:application/csv," + encodeURIComponent(CsvString);
 var x = document.createElement("A");
 x.setAttribute("href", CsvString );
 x.setAttribute("download","data_set.csv");
 document.body.appendChild(x);
 x.click();
}
</script>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ma1: [
  [1990, false],  [1991, false],  [1992, true],  [1993, false],
  [1994, false],  [1995, false],  [1996, false],  [1997, false],
  [1998, false],  [1999, false],  [2000, false],  [2001, false],
  [2002, false],  [2003, false],  [2004, false],  [2005, false],
  [2006, false],  [2007, false],  [2008, false],  [2009, false],
  [2010, false],  [2011, false],  [2012, false],  [2013, false],
  [2014, false],  [2015, false],  [2016, false]
];

ma2: [
  ["Palstond", false],
  ["Elamat", true],
  ["Unknown", false]
];

ta3: 1869.80;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ma1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ma2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* node 1: check answers correct */
node1_test: if ans1=1992 and ans2="Elamat" and ans3=1869.80 then 1 else 0;

/* node 2: fails to compare the units in each question */
node2_test: if ans1=2013 and ans2="Palstond" and ans3=186.98 then 1 else 0;

/* node 3: treats 2000-16 stations as unknown */
node3_test: if ans1=1992 and member(ans2, ["Palstond", "Unknown"]) and ans3=1869.80 then 1 else 0;

/* node 4: disregards the order of the columns */
node4_test: if ans1=2013 and ans2="Elamat" and member(ans3, [1869.8, 186.98]) then 1 else 0;

/* node 5: disregards the order of the columns */
node5_test: if ans1=1992 and ans2="Palstond" and member(ans3, [17567.7, 175677]) then 1 else 0;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check correct answers</description>
        <answertest>AlgEquiv</answertest>
        <sans>node1_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Great! It appears that when combining the data sets, you have taken into account the different units used!
    <br>In addition, it seems that you have correctly appended the data sets together.
</p>

<br>

<p>There is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>fails to compare the units in each question</description>
        <answertest>AlgEquiv</answertest>
        <sans>node2_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It appears that when combining the data sets, you have not taken into account the different units used.
    <br>Have a look at a plot of your rainfall by year. Does anything seem odd here?
    <br>We mentioned in the initial question that rainfall is measured in cm in one data set and mm in another.
</p>

<br>

<p style="margin-left:30px"><img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-1.jpg?raw=true"></p>
<br>
<p>Try again, but if stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>treats 2000-16 stations as unknown</description>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It appears that when combining the data sets, you have taken into account the different units used, which is great!
    <br>However, have a look at a plot of your rainfall by year with station included. Does anything seem odd here?
    <br>All stations from the year 2000 are unknown. Have a look at the station variable name in the two initial data sets. Can you see why this might have happened when combining the data?
</p>

<br>

<p style="margin-left:30px"><img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-2.jpg?raw=true"></p>
<br>
<p>Try again, but if stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>disregards the order of the columns</description>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It appears that when combining the data sets, you have not taken into account the order of the columns.
    <br>Have a look at the data - the columns are in a different order in the two data sets.
    <br>We can visualise this further by a plot of rainfall by year
    <br>We want to put both of the data sheets together. For this we go to Prepare &gt; Data Reshape &gt; Append dialogue in R-Instat. We suggest having a look at the tutorial on merging and appending data in R-Instat.
</p>

<br>

<p style="margin-left:30px"><img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-4.jpg?raw=true"></p>
<br>
<p>Try again, but if stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>disregards the order of the columns</description>
        <answertest>AlgEquiv</answertest>
        <sans>node5_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It appears that when combining the data sets, you have not taken into account the order of the columns.
    <br>Have a look at the data - the columns are in a different order in the two data sets.
    <br>We can visualise this further by a plot of rainfall by year
    <br>We want to put both of the data sheets together. For this we go to Prepare &gt; Data Reshape &gt; Append dialogue in R-Instat. We suggest having a look at the tutorial on merging and appending data in R-Instat.
</p>

<br>

<p style="margin-left:30px"><img src="https://github.com/IDEMSInternational/wxmaxima-stack/blob/main/data-files/appending-data-3.jpg?raw=true"></p>

<p>Try again, but if stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite - are you looking at the <b>overall rainfall</b> for each year?
    <br>We want to put both of the data sheets together. For this we go to Prepare &gt; Data Reshape &gt; Append dialogue in R-Instat. We suggest having a look at the tutorial on merging and appending data in R-Instat.
</p>
<br>
<p>Try again, but if stuck there is a worked solution of this problem <a href="https://ecampus.idems.international/course/view.php?id=238&amp;section=97">here</a>.
</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
